# Data structures

Once we start working with many pieces of data at once, it will be convenient for us to store data in structures like arrays or dictionaries (rather than just relying on variables).<br>

Types of data structures covered:
1. Tuples
2. Dictionaries
3. Arrays

<br>
As an overview, tuples and arrays are both ordered sequences of elements (so we can index into them). Dictionaries and arrays are both mutable.
We'll explain this more below!

## Tuples

We can create a tuple by enclosing an ordered collection of elements in `( )`.

Syntax: <br>
```julia
(item1, item2, ...)
```


```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

Output:
```
("penguins", "cats", "sugargliders")
```

We can index into this tuple,

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

but since tuples are immutable, we can't update it

```julia
myfavoriteanimals[1] = "otters"
```

Output:
```
Stacktrace:

 [1] top-level scope

   @ c:\Users\joaqu\Documents\Julia\Course_Introduction_to_Julia\2 - Data structures.ipynb:1
```

## Now in 1.6: NamedTuples

As you might guess, `NamedTuple`s are just like `Tuple`s except that each element additionally has a name! They have a special syntax using `=` inside a tuple:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Output:
```
(bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Like regular `Tuples`, `NamedTuples` are ordered, so we can retrieve their elements via indexing:

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

They also add the special ability to access values by their name:

```julia
myfavoriteanimals.bird
```

Output:
```
"penguins"
```

## Dictionaries

If we have sets of data related to one another, we may choose to store that data in a dictionary. We can create a dictionary using the `Dict()` function, which we can initialize as an empty dictionary or one storing key, value pairs.

Syntax:
```julia
Dict(key1 => value1, key2 => value2, ...)
```

A good example is a contacts list, where we associate names with phone numbers.

```julia
myphonebook = Dict("Jenny" => "867-5309", "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String, String} with 2 entries:
  "Jenny"        => "867-5309"
  "Ghostbusters" => "555-2368"
```

In this example, each name and number is a "key" and "value" pair. We can grab Jenny's number (a value) using the associated key

```julia
myphonebook["Jenny"]
```

Output:
```
"867-5309"
```

We can add another entry to this dictionary as follows

```julia
myphonebook["Kramer"] = "555-FILK"
```

Output:
```
"555-FILK"
```

Let's check what our phonebook looks like now...

```julia
myphonebook
```

Output:
```
Dict{String, String} with 3 entries:
  "Jenny"        => "867-5309"
  "Kramer"       => "555-FILK"
  "Ghostbusters" => "555-2368"
```

We can delete Kramer from our contact list - and simultaneously grab his number - by using `pop!`

```julia
pop!(myphonebook, "Kramer")
```

Output:
```
"555-FILK"
```

```julia
myphonebook
```

Output:
```
Dict{String, String} with 2 entries:
  "Jenny"        => "867-5309"
  "Ghostbusters" => "555-2368"
```

Unlike tuples and arrays, dictionaries are not ordered. So, we can't index into them.

```julia
myphonebook[1]
```

Output:
```
@ Base .\dict.jl:484

 [2] top-level scope

   @ c:\Users\joaqu\Documents\Julia\Course_Introduction_to_Julia\2 - Data structures.ipynb:1
```

In the example above, `julia` thinks you're trying to access a value associated with the key `1`.

## Vectors (Arrays)

Unlike tuples, arrays are mutable. Unlike dictionaries, arrays contain ordered collections. <br>
We can create an array by enclosing this collection in `[ ]`.

Syntax: <br>
```julia
[item1, item2, ...]
```


For example, we might create an array to keep track of my friends

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]
```

Output:
```
5-element Vector{String}:
 "Ted"
 "Robyn"
 "Barney"
 "Lily"
 "Marshall"
```

or to store a sequence of numbers

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mixture = [1, 1, 2, 3, "Ted", "Robyn"]
```

Output:
```
6-element Vector{Any}:
 1
 1
 2
 3
  "Ted"
  "Robyn"
```

Once we have an array, we can grab individual pieces of data from inside that array by indexing into the array. For example, if we want the third friend listed in `myfriends`, we write

```julia
myfriends[3]
```

Output:
```
"Barney"
```

We can use indexing to edit an existing element of an array

```julia
myfriends[3] = "Baby Bop"
```

Output:
```
"Baby Bop"
```

Yes, **Julia is 1-based indexing, not 0-based like Python.**  Wars are fought over lesser issues. I have a friend with the wisdom of Solomon who proposes settling this once and for all with ½ 😃

We can also edit the array by using the `push!` and `pop!` functions. `push!` adds an element to the end of an array and `pop!` removes the last element of an array.

We can add another number to our fibonnaci sequence

```julia
push!(fibonacci, 21)
println(fibonacci)
```

Output:
```
[1, 1, 2, 3, 5, 8, 13, 21]
```

and then remove it

```julia
pop!(fibonacci)
println(fibonacci)
```

Output:
```
[1, 1, 2, 3, 5, 8, 13]
```

So far I've given examples of only 1D arrays of scalars, but arrays can have an arbitrary number of dimensions and can also store other arrays. 
<br><br>
For example, the following are arrays of arrays:

```julia
favorites = [["koobideh", "chocolate", "eggs"],["penguins", "cats", "sugargliders"]]
```

Output:
```
2-element Vector{Vector{String}}:
 ["koobideh", "chocolate", "eggs"]
 ["penguins", "cats", "sugargliders"]
```

```julia
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Output:
```
3-element Vector{Vector{Int64}}:
 [1, 2, 3]
 [4, 5]
 [6, 7, 8, 9]
```

```julia
numbers[1][2]
```

Output:
```
2
```

**Be careful when you want to copy arrays!**

```julia
println(fibonacci)
```

Output:
```
[1, 1, 2, 3, 5, 8, 13]
```

```julia
somenumbers = fibonacci
somenumbers[1] = 404
println(fibonacci)
```

Output:
```
[404, 1, 2, 3, 5, 8, 13]
```

Editing `somenumbers` caused `fibonacci` to get updated as well!

In the above example, we didn't actually make a copy of `fibonacci`. We just created a new way to access the entries in the array bound to `fibonacci`.

If we'd like to make a copy of the array bound to `fibonacci`, we can use the `copy` function.

```julia
# First, restore fibonacci
fibonacci[1] = 1

somemorenumbers = copy(fibonacci)

somemorenumbers[1] = 404
println(fibonacci)
```

Output:
```
[1, 1, 2, 3, 5, 8, 13]
```

In this last example, fibonacci was not updated. Therefore we see that the arrays bound to `somemorenumbers` and `fibonacci` are distinct.

## Matrixes

We can create a random matrix with rand:

```julia
rand(4, 3)
```

Output:
```
4×3 Matrix{Float64}:
 0.106544   0.658299  0.363108
 0.378608   0.919002  0.874921
 0.905034   0.461783  0.407715
 0.0294152  0.973284  0.827029
```

```julia
rand(4, 3, 2)
```

Output:
```
4×3×2 Array{Float64, 3}:
[:, :, 1] =
 0.37878   0.271004  0.849516
 0.932181  0.405798  0.556773
 0.970954  0.859176  0.546097
 0.299711  0.8579    0.457414

[:, :, 2] =
 0.65714    0.606199  0.69653
 0.0955298  0.357481  0.624067
 0.541801   0.198771  0.312223
 0.871156   0.383433  0.347434
```

```julia
M = rand(3,3)
```

Output:
```
3×3 Matrix{Float64}:
 0.135991  0.212322  0.0610028
 0.917255  0.281517  0.631567
 0.123867  0.915488  0.837812
```

```julia
M[1,2]
```

Output:
```
0.21232192998558497
```

### Exercises

#### 3.1 
Create an array, `a_ray`, with the following code:

```julia
a_ray = [1, 2, 3]
```

Add the number `4` to the end of this array and then remove it.

```julia
a_ray = [1,2,3]
push!(a_ray, 4)
pop!(a_ray)
```

Output:
```
4
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Try to add "Emergency" as key to `myphonebook` with the value `string(911)` with the following code
```julia
myphonebook["Emergency"] = 911
```

Why doesn't this work?

```julia
myphonebook["Emergency"] = 911
```

Output:
```
@ Base .\dict.jl:369

 [2] top-level scope

   @ c:\Users\joaqu\Documents\Julia\Course_Introduction_to_Julia\2 - Data structures.ipynb:1
```

#### 3.3 
Create a new dictionary called `flexible_phonebook` that has Jenny's number stored as an integer and Ghostbusters' number stored as a string with the following code

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String, Any} with 2 entries:
  "Jenny"        => 8675309
  "Ghostbusters" => "555-2368"
```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 3.4 
Add the key "Emergency" with the value `911` (an integer) to `flexible_phonebook`.

```julia
flexible_phonebook["Emergency"] = 911
```

Output:
```
911
```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
haskey(flexible_phonebook, "Emergency")
```

Output:
```
true
```

#### 3.5 
Why can we add an integer as a value to `flexible_phonebook` but not `myphonebook`? How could we have initialized `myphonebook` so that it would accept integers as values? (hint: try using [Julia's documentation for dictionaries](https://docs.julialang.org/en/v1/base/collections/#Dictionaries))

*Answer:* Dicts can be created by passing pair objects constructed with => to a Dict constructor: Dict("A"=>1, "B"=>2). This call will attempt to infer type information from the keys and values (i.e. this example creates a Dict{String, Int64}). To explicitly specify types use the syntax Dict{KeyType,ValueType}(...). For example, Dict{String,Int32}("A"=>1, "B"=>2).

```julia
a = Dict{String, Any}
```

Output:
```
Dict{String, Any}
```